Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "ADEWALE ADEBOGUN ODUNSI"
COLLABORATORS = "2110481"
```

---

```python

```

---
# Welcome to CE880
### This is your week-1 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function that calculate the area of the square by using the formula as `area = side*side`
At last, print the area of the square to get the output.

<!-- grade_id: areaofSqr_ans, points: 0, grade: false, solution: true -->
```python
def areaofSqr(side):
    """Compute the area of squares with side = any number.
    """
    # YOUR CODE HERE
    area= side*side
    return area
    raise NotImplementedError()
```

```python
areaofSqr(10)
```

Output:
```
100
```

<!-- grade_id: areaofSqr_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert areaofSqr(1) == 1
assert areaofSqr(2) == 4
assert areaofSqr(-3) == 9
assert areaofSqr(10) == 100
assert areaofSqr(11) == 121
```

---

## Question 2:

A `trapezoid` is a quadrilateral with two sides parallel. The `trapezoid` is equivalent to the British definition of the `trapezium`. An isosceles `trapezoid` is a `trapezoid` in which the base angles are equal so

<img src="https://www.w3resource.com/w3r_images/python-math-image-exercise-3.svg"
     alt="Markdown Monster icon"
     style="float: left; margin-right: 10px;" />

<!-- grade_id: areaofTrap_ans, points: 0, grade: false, solution: true -->
```python
def areaofTrap(base1, base2, height):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    area= 0.5 * (base1 + base2) * height
    return area
    raise NotImplementedError()
```

```python
areaofTrap(5, 3, -9)
```

Output:
```
-36.0
```

<!-- grade_id: areaofTrap_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert areaofTrap(10, 10, 20) == 200.0
assert areaofTrap(5, 10, 20) == 150.0
assert areaofTrap(5, 3, 5) == 20.0
assert areaofTrap(5, 3, -9) == -36.0
```

---
## Question 3: 

 Write a Python program to reverse a string. Example: Sample String: `1234xyz` Expected Output: `zyx4321`

<!-- grade_id: reversestring_ans, points: 0, grade: false, solution: true -->
```python
def reverse_string(str1):
    """Write your code here to reverse the string"""
    
    # YOUR CODE HERE
    return str1[::-1]
    
    raise NotImplementedError()
```

```python
reverse_string('CE880')
```

Output:
```
'088EC'
```

<!-- grade_id: reversestring_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert reverse_string('Colchester') == 'retsehcloC'
assert reverse_string('University') == 'ytisrevinU'
assert reverse_string('Data Science') == 'ecneicS ataD'
assert reverse_string('CE880') == '088EC'
```

## Question 4:
Write a Python program to print the even numbers from a given list.

Sample List : `[1, 2, 3, 4, 6, 7, 8, 9]` Expected Result : `[2, 4, 6, 8]`

<!-- grade_id: even_num_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def even_num(lst_1):
    """Write a Python program to print the even numbers from a given list."""
    
    # YOUR CODE HERE  
    lst_1.sort()
    lst_1 = list(dict.fromkeys(lst_1))
    return list(filter(lambda num: (num % 2 == 0), lst_1))
    
    raise NotImplementedError()
```

```python
even_num([2, 3, 14, 17, 8, 19, 11, 14])
```

Output:
```
[2, 8, 14]
```

<!-- grade_id: even_num_test, points: 1, grade: true, solution: false -->
```python
even_num([1, 2, 3, 4, 6, 7, 8, 9]) == [2, 4, 6, 8]
even_num([2, 3, 4, 7, 8, 9, 11, 14]) == [2, 4, 8, 14]
even_num([2, 3, 14, 17, 8, 19, 11, 14]) == [ 2,  8, 14]
```

Output:
```
True
```

## Question 5:
Write a Python function to create and print a list where the values are square of numbers between `1` and `n`, where `n` is user defined number: Example `sqr_values(4) == [1, 4, 9]`

<!-- grade_id: sqr_values_ans, points: 0, grade: false, solution: true -->
```python
def sqr_values(n):
    """Write a Python program to print the square of numbers between 1 and `n`"""
    
    # YOUR CODE HERE
    sqr_list = []
    for i in range(1, n):      
      sqr_list.append(i**2)
    return sqr_list
    raise NotImplementedError()
```

```python
sqr_values(4)
```

Output:
```
[1, 4, 9]
```

<!-- grade_id: sqr_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that sqr_values returns the correct answer for various inputs."""
assert sqr_values(4) == [1, 4, 9]
assert sqr_values(8) == [1, 4, 9, 16, 25, 36, 49]
assert sqr_values(9) == [1, 4, 9, 16, 25, 36, 49, 64]
assert sqr_values(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
```

## Question 6: 
Write a Python program to generate all permutations of a list in Python.

In mathematics, the notion of permutation relates to the act of arranging all the members of a set into some sequence or order, or if the set is already ordered, rearranging (reordering) its elements, a process called permuting. These differ from combinations, which are selections of some members of a set where order is disregarded.

<!-- grade_id: cal_days_ans, points: 0, grade: false, solution: true -->
```python
import itertools

def perm_list(lst):
    """Write a Python program to calculate number of days between two dates."""
    # YOUR CODE HERE
    
    permut_list = []  
    list_of_permutations = itertools.permutations(lst)
    #list_of_permutations = permutations([1, 2, 3])
    for i in list(list_of_permutations):
      permut_list.append(i)
    return permut_list
    

    raise NotImplementedError()
```

```python
perm_list([1,2, 4])
```

Output:
```
[(1, 2, 4), (1, 4, 2), (2, 1, 4), (2, 4, 1), (4, 1, 2), (4, 2, 1)]
```

<!-- grade_id: cal_days_test, points: 1, grade: true, solution: false -->
```python
assert perm_list([1,2]) == [(1, 2), (2, 1)]
assert perm_list([1,2, 4]) == [(1, 2, 4), (1, 4, 2), (2, 1, 4), (2, 4, 1), (4, 1, 2), (4, 2, 1)]
```

```python

```